---
id: 6920e90809062d7475efe005
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

Now that you know how `Math.random()` creates a floating point number and how `Math.floor()` rounds a value down to the nearest whole integer, you can combine them to generate a random whole number between two values.

Here is the formula for generating a random number between two values:

```js
Math.random() * (maximum - minimum) + minimum;
```

By wrapping that expression inside `Math.floor()`, you turn the result into a whole number.

Create a variable called `randomInt` and assign it the result of rounding the random number generated between your existing `min` and `max` values.

Then, log the `randomInt` variable to the console and refresh the page a few times to see different whole number results.

# --hints--

You should have a `randomInt` variable.

```js
assert.isNotNull(randomInt);
```

You should use the formula generate a random number between the `min` and `max` values provided above.

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

// Must use Math.random()
assert.match(code, /Math\.random\s*\(\s*\)/);

// Must use a range expression involving max and min
assert.match(code, /\(\s*max\s*-\s*min\s*\)/);

// Must add min at the end
assert.match(code, /\+\s*min/);
```

You should not hardcode the value for the `randomInt` variable.

```js
assert.notMatch(code, /randomInt\s*=\s*\d+/);
```

The `randomInt` variable should be a whole number. Make sure to use `Math.floor()`.

```js
assert.equal(Number.isInteger(randomInt), true);
```

You should log the `randomInt` variable to the console.

```js
assert.match(code, /console\.log\(\s*randomInt\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = 'MathBot';
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log(
  'The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.'
);

const randomNum = Math.random();
console.log(randomNum);

console.log(
  'The Math.floor() method rounds the value down to the nearest whole integer.'
);
const num = 10.99;
Math.floor(num);

console.log('Now, generate a random integer between two values.');

const min = 1;
const max = 50;

console.log(
  'Use Math.floor() and Math.random() together to generate random integer.'
);
--fcc - editable - region--;

--fcc - editable - region--;
```
