---
id: 6920e90809062d7475efe005
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

Now that you have learned how `Math.random()` generates a floating point number and how `Math.floor()` rounds a number down, you can combine them to generate a random integer number between two values.

The formula to generate a random number between two values is the following:

```js
Math.random() * (maximum - minimum) + minimum;
```

And the method rounds the value down to the nearest whole integer.

```js
const price = 10.99;
Math.floor(price);
```

Using both `Math.floor()` and `Math.random()` will produce a result that is a random integer between two values.

Create a variable called `randomInteger` and assign it the result of rouding the generated random number between your existing `min` and `max` values.

Then, log the `randomInt` variable to the console and refresh the page a few times to see different whole-number results.

# --hints--

You should have a `randomInt` variable.

```js
assert.isNotNull(randomInt);
```

The `randomInt` variable should be a number. Make sure to use `Math.floor()`

```js
assert.equal(Number.isInteger(randomInt), true);
```

You should assign the result of generating a value between the `min` and `max` values. You should use the formulae provided above.

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);
```

You should not hardcode the value for the `randomInt` variable.

```js
assert.notMatch(code, /randomInt\s*=\s*\d+/);
```

You should log the `randomInt` variable to the console.

```js
assert.match(code, /console\.log\(\s*randomInt\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = 'MathBot';
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log(
  'The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.'
);

const randomNum = Math.random();
console.log(randomNum);

console.log(
  'The Math.floor() method rounds the value down to the nearest whole integer.'
);
const num = 10.99;
Math.floor(num); //10

console.log('Now, generate a random integer between two values.');

const min = 5;
const max = 15;

const randomInt = Math.random() * (max - min) + min;
console.log(randomInt);

--fcc - editable - region--;

--fcc - editable - region--;
```
